
# Assignment No 3b
###### *Sibt ul Hussain*
----
## Goal

Your goal in this part of assigment is to implement a Decision Tree Classifier for categorical variables.

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas.

## Submission Instructions
You are required to submit the original notebook file on the Slate (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.



### Decision Tree Classifier

Now in this assignment we will be implementing the Decision Classifier for both Continuous and Categorical attributes.

Decision tree can be built by using any of the following split criterias, namely:
 - Information Gain
 - Gini Index
 - CART 

However, you are required here to implement the decision tree with information gain as splitting criterion.

Remember in my code i am not looking for maximizing the information gain, instead i am looking for minimizing the split entropy. Recall,
$$Information Gain  = H(D) - H(D_Y,D_N)$$

Where,

$H(D)$ is the data set entroy and $H(D_Y,D_N)$ is split entropy. Since $H(D)$ is constant for the given dataset so maximizing the entropy is equal to minimizing the split entropy and that is what is being represented in my code outputs.

```python
%pylab inline
import scipy.stats
from collections import defaultdict  # default dictionary 
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

```python
#Cutomize the Matplotlib for beautiful plots...
# import dmStyle
# dmStyle.customize_mpl()

```

```python
import pandas as pd
import tools as t # set of tools for plotting, data splitting, etc..
```

```python
def getSplits(categories):
    '''
        function returns list of splits for the given list of categorical variables...
        
        Input:
        ------------
            categories: a list of unique categories...
        
        Return:
        ------------
            list of splits(tuples) for given list of categorical variables. Each pair of sublists
            defines the left and right splits, e.g. This list
            [('y', 'f'), ('s', 'g'), ('f'), ('y', 's', 'g')]
            
            defines two splits with each pair representing a different split.
        Examples:
        ------------
        splits=getSplits(['a1','a2','a3','a4']) will return 
        [('a1', 'a2', 'a4'), ('a3',), ('a2', 'a4'), 
        ('a1', 'a3'), ('a1', 'a4'), ('a3', 'a2'), ('a1', 'a3', 'a2'), 
        ('a4',), ('a3', 'a4'), ('a1', 'a2'), ('a1',), ('a3', 'a2', 'a4'), ('a2',), ('a1', 'a3', 'a4')]

            
    '''
    categories=set(categories)
    tsplits=t.get_powerset(categories,len(categories)-1) # get all the power sets with the given cardinality...
    flist=[]
    for s in tsplits:

        if not s in flist:
            r=categories.difference(s)
    #         print s, tuple(r)
            flist.append(s)
            flist.append(r)

    olist=[]

    for s in flist:
        ilist=[]

        for k in s:
            ilist.append(k)
        olist.append(tuple(ilist))

#     print olist
    
    return olist
```

```python
splits=getSplits(['a1','a2','a3','a4'])
print splits,len(splits)
```

Output:
```
[('a1', 'a2', 'a4'), ('a3',), ('a2', 'a4'), ('a1', 'a3'), ('a1', 'a4'), ('a3', 'a2'), ('a1', 'a3', 'a2'), ('a4',), ('a3', 'a4'), ('a1', 'a2'), ('a1',), ('a3', 'a2', 'a4'), ('a2',), ('a1', 'a3', 'a4')] 14
```

<!-- grade_id: node, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self,purity,klasslabel='',score=0,split=[],fidx=-1):
        '''
            purity: purity level at which to stop
            klasslabel: klasslabel of the node, (for leaf node)
            score: information gain of the newly added node
            split: splitting threshold
            fidx: feature index            
        '''
        self.lchild=None
        self.rchild=None
        self.klasslabel=klasslabel        
        self.split=split
        self.score=score
        self.fidx=fidx
        self.purity=purity
        self.ftype= 'categorical' if type(self.split) in [tuple, str, numpy.string_] else 'continuous' # feature type 

        
    def set_childs(self,lchild,rchild):
        # YOUR CODE HERE
        #raise NotImplementedError()
        self.lchild=lchild
        self.rchild=rchild
    def isleaf(self):
        # YOUR CODE HERE
        #raise NotImplementedError()
        if (self.rchild==None and self.lchild==None):
            return True
        else:
            return False
    def isless_than_eq(self, X):
        # YOUR CODE HERE
        #raise NotImplementedError()
        if(X[self.fidx]<=self.split):
            return True
        else:
            return False
    def get_str(self):        
        if self.isleaf():
            return 'C(class={},Purity={})'.format(self.klasslabel,self.purity)
        else:
            return 'I(Fidx={},Score={},Split={})'.format(self.fidx,self.score,self.split)
    
```

```python
# A placeholder class 
# TODO: You have to implement the following class, remember from the lectures that you will 
# need to build a model for each different class you are trying to identify...
```

<!-- grade_id: tree, points: 0, grade: false, solution: true -->
```python
### import pdb
## Your code goes here...
# You might need to define auxliary classes for composition.. ?
class DecisionTree:
    ''' Implements the Decision Tree For Classification... '''
    def __init__(self, purityp, exthreshold, maxdepth=10):        
        self.purity = purityp
        self.exthreshold = exthreshold
        self.maxdepth = maxdepth
        self.tree=None
        pass
    def train(self, X, Y):
        ''' Train Decision Tree using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        nexamples, nfeatures = X.shape
        # # now go and train a model for each class...
        # YOUR CODE HERE
        #raise NotImplementedError()
        self.tree=self.build_tree(X,Y,self.maxdepth)
        
    def build_tree(self, X, Y, depth):
        """ 
            Function is used to recursively build the decision Tree 
          
            Input
            -----
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns
            -------
            root node of the built tree...
        """
        nexamples, nfeatures = X.shape
        
        
        # YOUR CODE HERE
        #raise NotImplementedError()
        #return node
        klasses=np.unique(Y);
        #------------------------------------------------->>>>
        purities=[]
        # YOUR CODE HERE
        #raise NotImplementedError()
        #root=Node()
        for c in klasses:
            temp_sum=np.sum(Y==c)
            temp_purity=temp_sum/float(nexamples)
            purities.append(temp_purity)
        np_purities=np.array(purities)
        max_purity=np_purities[np.argmax(np_purities)]
        split_column=-1
        split_point_score=0
        major_class=''
        s=0
        ridx=[]
        lidx=[]
        if(max_purity>=self.purity or nexamples<=self.exthreshold or depth<=1):
            index=np.argmax(np_purities)
            major_class=klasses[index]
            my_node=Node(max_purity,major_class)
            return my_node
        else:
            #print "nexamples:",nexamples
            for i in range(0,nfeatures):
                if((len(np.unique(X[:,i])))>1):
                    split,mingain,Xlidx,Xridx=self.evaluate_categorical_attribute(X[:,i],Y)
                    if (split_point_score>mingain or split_column==-1):
                        split_point_score=mingain
                        split_column=i
                        s=split
                        lidx=Xlidx
                        ridx=Xridx
            if(len(lidx)>1 and len(ridx)>1):       
                index=np.argmax(np_purities)
                major_class=klasses[index] 
                my_node=Node(max_purity,major_class,split_point_score,s,split_column)
                #print 'X sahpe',lidx
                #print 'X sahpe',ridx
                leftchild=self.build_tree(X[lidx,:],Y[lidx],depth-1)
                rightchild=self.build_tree(X[ridx,:],Y[ridx],depth-1)
                my_node.set_childs(leftchild,rightchild)
                return my_node
            else:
                index=np.argmax(np_purities)
                major_class=klasses[index]
                my_node=Node(max_purity,major_class)
                return my_node
        
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        
        nexamples, nfeatures = X.shape
        pclasses = self.predict(X)
        
        # your code go here...
        
        return np.array(pclasses)
    def evaluate_categorical_attribute(self, feat, Y):
            """
            Evaluates the cateogrical attribute for all possible split points, i.e. 2^(m-1)-2 for
            possible feature selection

            Input:
            ---------
            feat: a categorical feature
            Y: labels

            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node

            """
            categories = set(feat)
            splits = getSplits(categories) if len(categories) > 1 else tuple(categories)
            freq = scipy.stats.itemfreq(Y)
            nclasses=np.unique(Y)
            ncategories=np.unique(feat)
            #print 'nclasses',nclasses
            #print 'ncategories',ncategories
            #index=np.argwhere(('a1','a2'))
            t=np.asarray(('a1','a4'))
            index=ncategories.searchsorted(t)
            #print 'idhar dekh BC:',index[1]
            #print 'the index is',index[0]
            #print "ncategories:",ncategories
            #print "the val at index 0 is:",ncategories[index]
            #---------------------------------------------->
            my_table=np.zeros((len(ncategories),len(nclasses)))
            for i in range(0,len(feat)):
                class_num=-1
                cat_num=-1
                for j in range(0,len(nclasses)):
                    if(Y[i]==nclasses[j]):
                        class_num=j
                        break
                for k in range(0,len(ncategories)):
                    if(feat[i]==ncategories[k]):
                        cat_num=k
                        break
                my_table[cat_num][class_num]=my_table[cat_num][class_num]+1
            #print 'here it is:', my_table
            #pd.series()
            #print 'splits',splits
            if splits>1:
                split_scores=np.zeros((len(splits)/2))
            else:
                split_scores=np.zeros((len(splits)))
            split_count=0
            #+++++++++++++++++++++++++++++++++++++++++++++++
            min_gain=999
            #for s in range(0,(len(splits)/2)):
            s=0
            while(s<(len(splits))-1):
                temp_tuple=splits[s]
                temp_tuple2=splits[s+1]
                DY=0
                DN=0
                for f in feat:
                    if(f in temp_tuple):
                        DY=DY+1
                    else:
                        DN=DN+1
                DY_prob=(DY/float(DY+DN))
                DN_prob=(DN/float(DY+DN))
                #print "DY:",DY_prob
                #print "DN:",DN_prob
                index=ncategories.searchsorted(np.asarray(temp_tuple))
                #index=np.argwhere(np.asarray(temp_tuple))
                index=index.flatten()
                #index2=np.argwhere(temp_tuple2)
                index2=ncategories.searchsorted(np.asarray(temp_tuple2))
                index2=index2.flatten()
                #print 'index1 is',index
                #print 'index2 is',index2
                sum1=np.sum(my_table[index,:],axis=1)
                #print "sum1:",sum1
                sum2=np.sum(my_table[index2,:],axis=1)
                #print "sum2:",sum2
                sum3=np.sum(my_table[index,:],axis=0)
                #print "sum3:",sum3
                sum4=np.sum(my_table[index2,:],axis=0)
                #print "sum4:",sum4
                sum3=sum3[sum3>0]
                sum4=sum4[sum4>0]
                sum3=sum3/float(np.sum(sum1))
                sum4=sum4/float(np.sum(sum2))
                temp=sum3*np.log2(sum3)
                temp2=sum4*np.log2(sum4)
                HDY=-1*np.sum(temp)
                HDN=-1*np.sum(temp2)
                split_entropy=(DY_prob*HDY)+(DN_prob*HDN)
                #print 'idhar dekho bc:', split_entropy
                if(split_entropy<min_gain):
                    best_split=[temp_tuple,temp_tuple2]
                    min_gain=split_entropy
                    #print 'the new min split entropy is:',split_entropy
                split_scores[split_count]=split_entropy
                split_count=split_count+1
                #print 'does it work:',temp
                #print 'probs for Dy are:',sum3
                #print 'probs for Dn are:',sum4
                s=s+2
            #+++++++++++++++++++++++++++++++++++++++++++++++
            #---------------------------------------------->
            # YOUR CODE HERE
            #raise NotImplementedError()
            mingain=split_scores[np.argmin(split_scores)]
            #print "the min gain is:",mingain
            csplit=best_split[0]
            Xlidx=[]
            Xridx=[]
            for z in range(0,len(feat)):
                if(feat[z] in csplit[0]):
                    Xlidx.append(z)
                else:
                    Xridx.append(z)
            #print 'the best split is:',best_split
            #print "Xlidx:",Xlidx
            #print "Xridx:",Xridx
            #Xlidx=np.array(Xlidx)
            #Xridx=np.array(Xridx)
            return csplit, mingain, Xlidx, Xridx
        
    def evaluate_numerical_attribute(self, feat, Y):
        '''
            Evaluates the numerical attribute for all possible split points for
            possible feature selection
            
            Input:
            ---------
            feat: a contiuous feature
            Y: labels
            
            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node
            
        '''
        
        classes = np.unique(Y)
        nclasses = len(self.classes)
        f = np.sort(feat)
        
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        return split, mingain, Xlidx, Xridx
    def predict(self, X):
        
        """
        Test the trained classifiers on the given example X
        
                   
            Input:
            ------
            X: [1 x d] a d-dimensional test example.
           
            Returns:
            -----------
                pclass: the predicted class for the given example, i.e. to which it belongs
        """
        z = []
        
        for idx in range(X.shape[0]):
            
            z.append(self._predict(self.tree, X[idx, :]))
        
        return z 
    
    def _predict(self, node, X):
        # YOUR CODE HERE
        #raise NotImplementedError()
        if(len(node.split)==0):
            return node.klasslabel
        elif(X[node.fidx] in node.split[0] and node.lchild!=None):
            return self._predict(node.lchild,X)
        elif(node.rchild!=None):
            return self._predict(node.rchild,X)
        else:
            return node.klasslabel

    def __str__(self):
        str = '---------------------------------------------------'
        str += '\n A Decision Tree With Depth={}'.format(self.find_depth())
        str += self.__print(self.tree)
        str += '\n---------------------------------------------------'
        return str  # self.__print(self.tree)        
        
     
   
    def find_depth(self):
        return self._find_depth(self.tree)
    def _find_depth(self, node):
        if not node:
            return
        if node.isleaf():
            return 1
        else:
            return max(self._find_depth(node.lchild), self._find_depth(node.rchild)) + 1
    def __print(self, node, depth=0):
        
        ret = ""

        # Print right branch
        if node.rchild:
            ret += self.__print(node.rchild, depth + 1)

        # Print own value
        
        ret += "\n" + ("    "*depth) + node.get_str()

        # Print left branch
        if node.lchild:
            ret += self.__print(node.lchild, depth + 1)
        
        return ret
```

# Lets test our code for the given example in the book.

```python
#load the Iris dataset
tdata=pd.read_csv('./iris.data')
tdata.columns=['SepalLength','SepalWidth','PetalLength','PetalWidth','Class']
```

```python
tx=tdata['SepalLength'].dropna()
tx[(tdata['SepalLength']>=4.3) & (tdata['SepalLength']<=5.2)]='a1'
tx[(tdata['SepalLength']>5.2) & (tdata['SepalLength']<=6.1)]='a2'
tx[(tdata['SepalLength']>6.1) & (tdata['SepalLength']<=7.0)]='a3'
tx[(tdata['SepalLength']>7.0) & (tdata['SepalLength']<=7.9)]='a4'
```

```python
print tx.values
Y=tdata['Class'].dropna()
Y[Y=='Iris-virginica']='Iris-versicolor'
Y=Y.values
print Y
```

Output:
```
['a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a2' 'a2'
 'a2' 'a1' 'a2' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a2' 'a1' 'a2' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a1' 'a1' 'a2' 'a1' 'a3' 'a3' 'a3' 'a2' 'a3' 'a2' 'a3' 'a1' 'a3' 'a1' 'a1'
 'a2' 'a2' 'a2' 'a2' 'a3' 'a2' 'a2' 'a3' 'a2' 'a2' 'a2' 'a3' 'a2' 'a3' 'a3'
 'a3' 'a3' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a3' 'a3' 'a2' 'a2' 'a2'
 'a2' 'a2' 'a1' 'a2' 'a2' 'a2' 'a3' 'a1' 'a2' 'a3' 'a2' 'a4' 'a3' 'a3' 'a4'
 'a1' 'a4' 'a3' 'a4' 'a3' 'a3' 'a3' 'a2' 'a2' 'a3' 'a3' 'a4' 'a4' 'a2' 'a3'
 'a2' 'a4' 'a3' 'a3' 'a4' 'a3' 'a2' 'a3' 'a4' 'a4' 'a4' 'a3' 'a3' 'a2' 'a4'
 'a3' 'a3' 'a2' 'a3' 'a3' 'a3' 'a2' 'a3' 'a3' 'a3' 'a3' 'a3' 'a3' 'a2']
['Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'I
```

```python
%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx = dt.evaluate_categorical_attribute(tx,Y)
print 
print split, gain
```

Output:
```
Automatic pdb calling has been turned OFF

('a1',) 0.510702350316
```

<!-- grade_id: split_gain, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_almost_equals

%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx = dt.evaluate_categorical_attribute(tx,Y)


assert_almost_equal('a1', split[0])
assert_almost_equal(gain, 0.51, 1)
```

Output:
```
Automatic pdb calling has been turned OFF
```

# Now lets test on [Mushroom](https://archive.ics.uci.edu/ml/machine-learning-databases/mushroom/agaricus-lepiota.names) dataset


This data set includes descriptions of hypothetical samples corresponding to 23 species of gilled mushrooms in the Agaricus and Lepiota Family (pp. 500-525).  Each species is identified as definitely edible, definitely poisonous, or of unknown edibility and not recommended.  This latter class was combined with the poisonous one.  The Guide clearly states that there is no simple rule for determining the edibility of a mushroom; no rule like ``leaflets three, let it be'' for Poisonous Oak and Ivy.


- Number of Instances: 8124

- Number of Attributes: 22 (all nominally valued)

- Attribute Information: (classes: edible=e, poisonous=p)
         1. cap-shape:                bell=b,conical=c,convex=x,flat=f, knobbed=k,sunken=s
         2. cap-surface:              fibrous=f,grooves=g,scaly=y,smooth=s
         3. cap-color:                brown=n,buff=b,cinnamon=c,gray=g,green=r, pink=p,purple=u,red=e,white=w,yellow=y
         4. bruises?:                 bruises=t,no=f
         5. odor:                     almond=a,anise=l,creosote=c,fishy=y,foul=f, musty=m,none=n,pungent=p,spicy=s
         6. gill-attachment:          attached=a,descending=d,free=f,notched=n
         7. gill-spacing:             close=c,crowded=w,distant=d
         8. gill-size:                broad=b,narrow=n
         9. gill-color:               black=k,brown=n,buff=b,chocolate=h,gray=g, green=r,orange=o,pink=p,purple=u,red=e,white=w,yellow=y
        10. stalk-shape:              enlarging=e,tapering=t
        11. stalk-root:               bulbous=b,club=c,cup=u,equal=e, rhizomorphs=z,rooted=r,missing=?
        12. stalk-surface-above-ring: fibrous=f,scaly=y,silky=k,smooth=s
        13. stalk-surface-below-ring: fibrous=f,scaly=y,silky=k,smooth=s
        14. stalk-color-above-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        15. stalk-color-below-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        16. veil-type:                partial=p,universal=u
        17. veil-color:               brown=n,orange=o,white=w,yellow=y
        18. ring-number:              none=n,one=o,two=t
        19. ring-type:                cobwebby=c,evanescent=e,flaring=f,large=l, none=n,pendant=p,sheathing=s,zone=z
        20. spore-print-color:        black=k,brown=n,buff=b,chocolate=h,green=r, orange=o,purple=u,white=w,yellow=y
        21. population:               abundant=a,clustered=c,numerous=n, scattered=s,several=v,solitary=y
        22. habitat:                  grasses=g,leaves=l,meadows=m,paths=p, urban=u,waste=w,woods=d

- Missing Attribute Values: 2480 of them (denoted by "?"), all for  attribute #11.

- Class Distribution: 
        --    edible: 4208 (51.8%)
        -- poisonous: 3916 (48.2%)
        --     total: 8124 instances

```python
#load the data set
data=pd.read_csv('./mushrooms.csv')
data.columns=columns=['class', 'cap-shape','cap-surface','cap-color','bruises?','odor','gill-attachment','gill-spacing','gill-size','gill-color','stalk-shape','stalk-root','stalk-surface-above-ring','stalk-surface-below-ring','stalk-color-above-ring','stalk-color-below-ring','veil-type','veil-color','ring-number','ring-type','spore-print-color','population','habitat']
print data.describe()
```

Output:
```
class cap-shape cap-surface cap-color bruises?  odor gill-attachment  \
count   8124      8124        8124      8124     8124  8124            8124   
unique     2         6           4        10        2     9               2   
top        e         x           y         n        f     n               f   
freq    4208      3656        3244      2284     4748  3528            7914   

       gill-spacing gill-size gill-color   ...   stalk-surface-below-ring  \
count          8124      8124       8124   ...                       8124   
unique            2         2         12   ...                          4   
top               c         b          b   ...                          s   
freq           6812      5612       1728   ...                       4936   

       stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
count                    8124                   8124      8124       8124   
unique                      9                      9         1         
```

```python
data.tail()
```

Output:
```
class cap-shape cap-surface cap-color bruises? odor gill-attachment  \
8119     e         k           s         n        f    n               a   
8120     e         x           s         n        f    n               a   
8121     e         f           s         n        f    n               a   
8122     p         k           y         n        f    y               f   
8123     e         x           s         n        f    n               a   

     gill-spacing gill-size gill-color   ...   stalk-surface-below-ring  \
8119            c         b          y   ...                          s   
8120            c         b          y   ...                          s   
8121            c         b          n   ...                          s   
8122            c         n          b   ...                          k   
8123            c         b          y   ...                          s   

     stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
8119                  
```

```python
# Get your data in matrix
X=np.asarray(data.ix[:,1:].dropna())
Y=np.asarray(data.ix[:,0].dropna())
print " Data Set Dimensions=", X.shape, " True Class labels dimensions", Y.shape   

```

Output:
```
Data Set Dimensions= (8124, 22)  True Class labels dimensions (8124,)
```

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape   

```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
# Lets train a Decision Tree Classifier two features
#%pdb
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
```

```python
print dt
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=8
                            C(class=p,Purity=1.0)
                        I(Fidx=0,Score=0.992418479298,Split=('f',))
                            C(class=p,Purity=0.543263964951)
                    I(Fidx=1,Score=0.972249493155,Split=('s', 'g'))
                        C(class=p,Purity=0.647058823529)
                I(Fidx=1,Score=0.956110167195,Split=('f',))
                    C(class=e,Purity=0.677966101695)
            I(Fidx=0,Score=0.934317265819,Split=('b',))
                        C(class=e,Purity=0.838709677419)
                    I(Fidx=1,Score=0.605515196172,Split=('y', 'f'))
                        C(class=e,Purity=0.845454545455)
                I(Fidx=1,Score=0.457279513218,Split=('s',))
                    C(class=e,Purity=0.941860465116)
        I(Fidx=0,Score=0.94164538635,Split=('k', 'c', 'f'))
                    C(class=p,Purity=0.872807017544)
                I(Fidx=1,Score=0.7376
```

```python
pclasses=dt.predict(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
1513
Accuracy =  0.6208453016
```

<!-- grade_id: acc, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)


pclasses=dt.predict(Xtest[:,feat])
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.60)
```

<!-- grade_id: acc_all, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

dt=DecisionTree(0.95,5,10)
dt.train(Xtrain,Ytrain)

pclasses=dt.predict(Xtest)
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.90)
```

# Lets Train on all the features and for both the classes....

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape   

```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
feat=arange(22)#[0, 1, 2, 3]
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
pclasses=dt.predict(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
2404
Accuracy =  0.986458760771
```

```python
print dt
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=5
                C(class=p,Purity=1.0)
            I(Fidx=4,Score=0.108913334573,Split=('n',))
                C(class=e,Purity=0.964446260727)
        I(Fidx=4,Score=0.106265764785,Split=('c', 'f', 'm', 'p', 's', 'y'))
            C(class=p,Purity=1.0)
    I(Fidx=4,Score=0.103201836444,Split=('l', 'n'))
        C(class=e,Purity=1.0)
I(Fidx=4,Score=0.100131678225,Split=('a', 'l', 'n'))
    C(class=e,Purity=1.0)
---------------------------------------------------
```

What can you conclude ?
====================
Please write your observation....



```python
#Training on all features increases the accuracy because each feature distintly classifies each class
```

# Cross-Validation

Until now we have been splitting the dataset into a training and test set rather randomly and were reporting a rather artifical performance. Now we are going to test our system exhaustively by making use of k-fold [cross validation](http://en.wikipedia.org/wiki/Cross-validation_%28statistics%29). 

Now go and tune your hyper-parameters (purity, exthreshold) to opitmize the performance.

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=arange(22)#[0, 1, 2, 3] # features to use for our system
#now lets train and test on these folds...
totacc=[]
for k in range(nfolds):
    dt=DecisionTree(0.95,5,7)
    dt.train(folds[k][0][:,features],folds[k][1])
    pclasses=dt.predict(folds[k][2][:,features])
    acc=np.sum(pclasses==folds[k][3])/float(folds[k][3].shape[0])
    print dt
    print "[Info] Fold {} Accuracy = {}".format(k+1, acc)
    totacc.append(acc)

print totacc, '\n Mean Accuracy =', np.mean(totacc)
```

Output:
```
Generating CV data for 2 classes
---------------------------------------------------
 A Decision Tree With Depth=5
                C(class=p,Purity=1.0)
            I(Fidx=4,Score=0.109732043077,Split=('n',))
                C(class=e,Purity=0.964380447139)
        I(Fidx=4,Score=0.106789574939,Split=('c', 'f', 'm', 'p', 's', 'y'))
            C(class=p,Purity=1.0)
    I(Fidx=4,Score=0.103729148536,Split=('l', 'n'))
        C(class=e,Purity=1.0)
I(Fidx=4,Score=0.100792167171,Split=('a', 'l', 'n'))
    C(class=e,Purity=1.0)
---------------------------------------------------
[Info] Fold 1 Accuracy = 0.987198424421
---------------------------------------------------
 A Decision Tree With Depth=5
                C(class=p,Purity=1.0)
            I(Fidx=4,Score=0.109392222059,Split=('n',))
                C(class=e,Purity=0.964701464514)
        I(Fidx=4,Score=0.106551123665,Split=('c', 'f', 'm', 'p', 's', 'y'))
            C(class=p,Purity=1.0)
    I(Fidx=4,Score=0.103611124611,Split=('l'
```

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=[0,1, 2, 3] # features to use for our system
#now lets train and test on these folds...

#Lets perform the grid search...
purity=np.linspace(0.85,0.97,13)
nexamp=np.linspace(5,25,21)  

params=np.zeros((len(purity),len(nexamp)))
                   
for p in range(len(purity)):
    for n in range(len(nexamp)):
        totacc=[]
        for k in range(nfolds):
            dt=DecisionTree(purity[p],nexamp[n],7)#purity[p],nexamp[n])
            dt.train(folds[k][0][:,features],folds[k][1])
            pclasses=dt.predict(folds[k][2][:,features])
            acc=np.sum(pclasses==folds[k][3])/float(folds[k][3].shape[0])
            print "[Info] Fold {} Accuracy = {}".format(k+1, acc)
            totacc.append(acc)
        params[p,n]=np.mean(totacc)
        print totacc, '\nPurity={}, Nexample-threshold={}, Mean Accuracy ={}'.format(purity[p],nexamp[n], np.mean(totacc))
```

Output:
```
Generating CV data for 2 classes
[Info] Fold 1 Accuracy = 0.821762678484
[Info] Fold 2 Accuracy = 0.806991629739
[Info] Fold 3 Accuracy = 0.825701624815
[Info] Fold 4 Accuracy = 0.796159527326
[0.8217626784835057, 0.80699162973904481, 0.82570162481536191, 0.79615952732644013] 
Purity=0.85, Nexample-threshold=5.0, Mean Accuracy =0.812653865091
[Info] Fold 1 Accuracy = 0.821762678484
[Info] Fold 2 Accuracy = 0.806991629739
[Info] Fold 3 Accuracy = 0.825701624815
[Info] Fold 4 Accuracy = 0.796159527326
[0.8217626784835057, 0.80699162973904481, 0.82570162481536191, 0.79615952732644013] 
Purity=0.85, Nexample-threshold=6.0, Mean Accuracy =0.812653865091
[Info] Fold 1 Accuracy = 0.821762678484
[Info] Fold 2 Accuracy = 0.806991629739
[Info] Fold 3 Accuracy = 0.825701624815
[Info] Fold 4 Accuracy = 0.796159527326
[0.8217626784835057, 0.80699162973904481, 0.82570162481536191, 0.79615952732644013] 
Purity=0.85, Nexample-threshold=7.0, Mean Accuracy =0.812653865091
[Info] Fold 1 Accuracy = 0.8217
```

```python
np.unravel_index(np.argmax(params), params.shape)
```

Output:
```
(5, 18)
```

```python

```